# Assignment 2 - Network Connectivity

In this assignment you will go through the process of importing and analyzing an internal email communication network between employees of a mid-sized manufacturing company. 
Each node represents an employee and each directed edge between two nodes represents an individual email. The left node represents the sender and the right node represents the recipient. We will also store the timestamp of each email. 

```python
import networkx as nx
from collections import Counter
```

```python
graph = nx.read_edgelist(r"assets/email_network.txt", create_using = nx.MultiDiGraph(), nodetype = str, data = (("Recipient time", str), ),
                 edgetype = str)
```

### Question 1

Using networkx, load up the directed multigraph from `assets/email_network.txt`. Make sure the node names are strings.

*This function should return a directed multigraph networkx graph.*

```python
nx.get_edge_attributes(graph, "Recipient time");
```

<!-- grade_id: cell-60316968695c6f20, points: 0, grade: false, solution: true -->
```python
def answer_one():
    
    graph = nx.read_edgelist(r"assets/email_network.txt", create_using = nx.MultiDiGraph(), nodetype = str, data = (("Recipient time", str), ),
                 edgetype = str)
    return graph
```

<!-- grade_id: cell-483a56c89f9db231, points: 0, grade: true, solution: false -->
```python
ans_one = answer_one()
```

```python
ans_one.edges();
```

### Question 2

How many employees are represented in the network?

How many `sender`->`recepient` pairs of employees are there in the network such that `sender` sent at least one email to `recepient`? Note that even if a `sender` sent multiple messages to a `recepient`, they should only be counted once. You should exclude cases where an employee sent emails to themselves from this count. 

*This function should return a tuple with two integers (#employees, # `sender`->`recepient` pairs).*

<!-- grade_id: cell-16a78f22207d9f95, points: 0, grade: false, solution: true -->
```python
def answer_two():
    
    graph = answer_one()
    n_employees = len(set(graph.nodes()))
    sr_pairs = {pair for pair in graph.edges() if (pair[0] != pair[1])}    
    return(n_employees, len(sr_pairs))
```

<!-- grade_id: cell-5b6391549b076d2b, points: 1, grade: true, solution: false -->
```python
ans_two = answer_two()

```

```python
ans_two
```

Output:
```
(167, 5783)
```

### Question 3

* Part 1. Assume that information in this company can only be exchanged through email.

    When an employee sends an email to another employee, a communication channel has been created, allowing the sender to provide information to the reciever, but not viceversa. 

    Based on the emails sent in the data, is it possible for information to go from every employee to every other employee?


* Part 2. Now assume that a communication channel established by an email allows information to be exchanged both ways. 

    Based on the emails sent in the data, is it possible for information to go from every employee to every other employee?


*This function should return a tuple of bools (part1, part2).*

```python
nx.is_strongly_connected(graph)
```

Output:
```
False
```

```python
nx.is_weakly_connected(graph)
```

Output:
```
True
```

<!-- grade_id: cell-606dd37fb2f55358, points: 0, grade: false, solution: true -->
```python
def answer_three():
    graph = answer_one()
    return nx.is_strongly_connected(graph), nx.is_weakly_connected(graph)
```

<!-- grade_id: cell-82b3f0bc45e2895f, points: 1, grade: true, solution: false -->
```python
ans_three = answer_three()

```

```python
ans_three
```

Output:
```
(False, True)
```

### Question 4

How many nodes are in the largest weakly connected component of the graph?

*This function should return an int.*

```python
sorted([int(node) for node in list(nx.weakly_connected_components(graph))[0]])
```

Output:
```
[1,
 2,
 3,
 4,
 5,
 6,
 7,
 8,
 9,
 10,
 11,
 12,
 13,
 14,
 15,
 16,
 17,
 18,
 19,
 20,
 21,
 22,
 23,
 24,
 25,
 26,
 27,
 28,
 29,
 30,
 31,
 32,
 33,
 34,
 35,
 36,
 37,
 38,
 39,
 40,
 41,
 42,
 43,
 44,
 45,
 46,
 47,
 48,
 49,
 50,
 51,
 52,
 53,
 54,
 55,
 56,
 57,
 58,
 59,
 60,
 61,
 62,
 63,
 64,
 65,
 66,
 67,
 68,
 69,
 70,
 71,
 72,
 73,
 74,
 75,
 76,
 77,
 78,
 79,
 80,
 81,
 82,
 83,
 84,
 85,
 86,
 87,
 88,
 89,
 90,
 91,
 92,
 93,
 94,
 95,
 96,
 97,
 98,
 99,
 100,
 101,
 102,
 103,
 104,
 105,
 106,
 107,
 108,
 109,
 110,
 111,
 112,
 113,
 114,
 115,
 116,
 117,
 118,
 119,
 120,
 121,
 122,
 123,
 124,
 125,
 126,
 127,
 128,
 129,
 130,
 131,
 132,
 133,
 134,
 135,
 136,
 137,
 138,
 139,
 140,
 141,
 142,
 143,
 144,
 145,
 146,
 147,
 148,
 149,
 150,
 151,
 152,
 153,
 154,
 155,
 156,
 157,
 158,
 159,
 160,
 161,
 162,
 163,
 164,
 165,
 166,
 167]
```

<!-- grade_id: cell-39196f4c27b35534, points: 0, grade: false, solution: true -->
```python
def answer_four():
    graph = answer_one()
    return len(list(nx.weakly_connected_components(graph))[0])
```

<!-- grade_id: cell-2b1b7b06ecfa751d, points: 1, grade: true, solution: false -->
```python
ans_four = answer_four()
```

```python
ans_four
```

Output:
```
167
```

### Question 5

How many nodes are in the largest strongly connected component?

*This function should return an int*

```python
max([len(component) for component in nx.strongly_connected_components(graph)])
```

Output:
```
126
```

<!-- grade_id: cell-879d17a9db5c3790, points: 0, grade: false, solution: true -->
```python
def answer_five():
    graph = answer_one()
    return max([len(component) for component in nx.strongly_connected_components(graph)])
```

<!-- grade_id: cell-b0524f7dc1fbdec4, points: 1, grade: true, solution: false -->
```python
ans_five = answer_five()

```

```python
ans_five
```

Output:
```
126
```

### Question 6

Using the NetworkX functions `strongly_connected_components` and `subgraph`, find the subgraph of nodes in the largest strongly connected component. 
Call this graph G_sc.

*This function should return a networkx MultiDiGraph named G_sc.*

```python
largest_ccomp = sorted(nx.strongly_connected_components(graph), key = len)[-1]
```

```python
nx.subgraph(graph, largest_ccomp)
```

Output:
```
<networkx.classes.multidigraph.MultiDiGraph at 0x183e07a3d10>
```

<!-- grade_id: cell-a81e6f88ef84578c, points: 0, grade: false, solution: true -->
```python
def answer_six():
    graph = answer_one()
    largest_ccomp = sorted(nx.strongly_connected_components(graph), key = len)[-1]
    G_sc = nx.subgraph(graph, largest_ccomp)
    return G_sc
```

<!-- grade_id: cell-cf148ef273b3b19c, points: 0, grade: true, solution: false -->
```python
ans_six = answer_six()
assert type(ans_six) == nx.MultiDiGraph , "Your return type should be a MultiDiGraph object"

```

```python
len(ans_six.nodes())
```

Output:
```
126
```

### Question 7

What is the average distance between nodes in G_sc?

*This function should return a float.*

```python
nx.average_shortest_path_length(ans_six)
```

Output:
```
1.6461587301587302
```

<!-- grade_id: cell-1a3075039b3c2c0b, points: 0, grade: false, solution: true -->
```python
def answer_seven():
    graph = answer_six()
    return nx.average_shortest_path_length(graph)
```

<!-- grade_id: cell-5b374fdd48f37e02, points: 1, grade: true, solution: false -->
```python
ans_seven = answer_seven()

```

### Question 8

What is the largest possible distance between two employees in G_sc?

*This function should return an int.*

<!-- grade_id: cell-4fbb8d34d00f675c, points: 0, grade: false, solution: true -->
```python
def answer_eight():
    graph = answer_six()
    return nx.diameter(graph)
```

<!-- grade_id: cell-c5714787854ef644, points: 1, grade: true, solution: false -->
```python
ans_eight = answer_eight()

```

```python
ans_eight
```

Output:
```
3
```

### Question 9

What is the set of nodes in G_sc with eccentricity equal to the diameter?

*This function should return a set of the node(s).*

<!-- grade_id: cell-bcfab850744277ed, points: 0, grade: false, solution: true -->
```python
def answer_nine():
    graph = answer_six()
    diam = nx.diameter(graph)
    return {node for node in graph.nodes() if (nx.eccentricity(graph, node) == diam)}
```

<!-- grade_id: cell-77c9ca0b94df3d6f, points: 1, grade: true, solution: false -->
```python
ans_nine = answer_nine()
assert type(ans_nine) == set, "Student answer must return a set"
```

```python
ans_nine
```

Output:
```
{'129', '134', '97'}
```

### Question 10

What is the set of node(s) in G_sc with eccentricity equal to the radius?

*This function should return a set of the node(s).*

<!-- grade_id: cell-fd3d43ea37da03f7, points: 0, grade: false, solution: true -->
```python
def answer_ten():
    graph = answer_six()
    rad = nx.radius(graph)
    return {node for node in graph.nodes() if (nx.eccentricity(graph, node) == rad)}
```

<!-- grade_id: cell-bfd2ee304bc25264, points: 1, grade: true, solution: false -->
```python
ans_ten = answer_ten()
assert type(ans_ten) == set, "Student answer must return a set"
```

```python
ans_ten
```

Output:
```
{'38'}
```

### Question 11

Which node in G_sc has the most shortest paths to other nodes whose distance equal the diameter of G_sc?


For the node with the most such shortest paths, how many of these paths are there?


*This function should return a tuple (name of node, number of paths).*

```python
# consider using this
from collections import Counter
plens = list(nx.all_pairs_shortest_path_length(answer_six()))
diam = nx.diameter(answer_six())
[(node, Counter(path_len_dict.values())[diam]) for (node, path_len_dict) in plens]
```

Output:
```
[('1', 0),
 ('2', 0),
 ('3', 0),
 ('4', 0),
 ('5', 0),
 ('6', 0),
 ('7', 0),
 ('8', 0),
 ('9', 0),
 ('10', 0),
 ('11', 0),
 ('12', 0),
 ('13', 0),
 ('14', 0),
 ('15', 0),
 ('16', 0),
 ('17', 0),
 ('18', 0),
 ('19', 0),
 ('20', 0),
 ('21', 0),
 ('22', 0),
 ('23', 0),
 ('24', 0),
 ('25', 0),
 ('26', 0),
 ('27', 0),
 ('28', 0),
 ('29', 0),
 ('30', 0),
 ('31', 0),
 ('32', 0),
 ('33', 0),
 ('34', 0),
 ('35', 0),
 ('36', 0),
 ('37', 0),
 ('38', 0),
 ('39', 0),
 ('40', 0),
 ('41', 0),
 ('42', 0),
 ('43', 0),
 ('44', 0),
 ('45', 0),
 ('46', 0),
 ('47', 0),
 ('48', 0),
 ('49', 0),
 ('50', 0),
 ('51', 0),
 ('52', 0),
 ('53', 0),
 ('54', 0),
 ('55', 0),
 ('56', 0),
 ('57', 0),
 ('58', 0),
 ('59', 0),
 ('60', 0),
 ('61', 0),
 ('62', 0),
 ('63', 0),
 ('64', 0),
 ('65', 0),
 ('66', 0),
 ('67', 0),
 ('68', 0),
 ('69', 0),
 ('70', 0),
 ('71', 0),
 ('72', 0),
 ('73', 0),
 ('74', 0),
 ('75', 0),
 ('76', 0),
 ('77', 0),
 ('78', 0),
 ('79', 0),
 ('80', 0),
 ('81', 0),
 ('82', 0),
 ('83', 0),
 ('84', 0),
 
```

```python
list(nx.all_shortest_paths(graph, "10", "101"))
```

Output:
```
[['10', '54', '101'],
 ['10', '1', '101'],
 ['10', '45', '101'],
 ['10', '37', '101'],
 ['10', '75', '101'],
 ['10', '27', '101'],
 ['10', '6', '101'],
 ['10', '109', '101'],
 ['10', '46', '101'],
 ['10', '93', '101'],
 ['10', '18', '101'],
 ['10', '55', '101'],
 ['10', '19', '101'],
 ['10', '35', '101'],
 ['10', '57', '101'],
 ['10', '88', '101'],
 ['10', '84', '101'],
 ['10', '50', '101'],
 ['10', '4', '101']]
```

```python
not None
```

Output:
```
True
```

<!-- grade_id: cell-c40858baf3d92139, points: 0, grade: false, solution: true -->
```python
def answer_eleven():
    from collections import Counter
    graph = answer_six()
    diam = nx.diameter(graph)    
    
    plens = list(nx.all_pairs_shortest_path_length(graph))
    diam = nx.diameter(graph)
    return sorted([(node, Counter(path_len_dict.values())[diam]) for (node, path_len_dict) in plens],
                  key = lambda _tuple: _tuple[1], reverse = True)[0]
```

```python
answer_eleven()
```

Output:
```
('97', 63)
```

<!-- grade_id: cell-f79b06650f61cf37, points: 1, grade: true, solution: false -->
```python
ans_eleven = answer_eleven()
assert type(ans_eleven) == tuple, "Student answer must be a tuple"
```

### Question 12

Suppose you want to prevent communication flow from the node that you found in question 11 to node 10. What is the smallest number of nodes you would need to remove from the graph (you're not allowed to remove the node from the previous question or 10)? 

*This function should return an integer.*

```python
list(nx.all_shortest_paths(graph, "97", "10"))
```

Output:
```
[['97', '14', '10']]
```

<!-- grade_id: cell-a4606db8971dcbba, points: 0, grade: false, solution: true -->
```python
def answer_twelve():
   # nodes 97 -> node 10
    graph = answer_six()
    return len(nx.minimum_node_cut(graph, s = "97", t = "10"))
```

```python
answer_twelve()
```

Output:
```
1
```

<!-- grade_id: cell-509cfa9f4136124d, points: 1, grade: true, solution: false -->
```python
ans_twelve = answer_twelve()

```

### Question 13

Convert the graph G_sc into an undirected graph by removing the direction of the edges of G_sc. Call the new graph G_un. 


*This function should return a networkx Graph.*

<!-- grade_id: cell-612b40425f0851e8, points: 0, grade: false, solution: true -->
```python
def answer_thirteen():
    graph = answer_six()
    G_un = nx.to_networkx_graph(graph)
    return G_un
```

```python
answer_thirteen().nodes()
```

Output:
```
NodeView(('1', '2', '3', '4', '5', '6', '7', '8', '9', '10', '11', '12', '13', '14', '15', '16', '17', '18', '19', '20', '21', '22', '23', '24', '25', '26', '27', '28', '29', '30', '31', '32', '33', '34', '35', '36', '37', '38', '39', '40', '41', '42', '43', '44', '45', '46', '47', '48', '49', '50', '51', '52', '53', '54', '55', '56', '57', '58', '59', '60', '61', '62', '63', '64', '65', '66', '67', '68', '69', '70', '71', '72', '73', '74', '75', '76', '77', '78', '79', '80', '81', '82', '83', '84', '85', '86', '87', '88', '89', '90', '91', '92', '93', '94', '95', '96', '97', '98', '99', '100', '101', '103', '104', '105', '106', '107', '108', '109', '110', '111', '112', '113', '115', '117', '118', '119', '120', '121', '122', '123', '124', '128', '129', '134', '149', '151'))
```

<!-- grade_id: cell-d1c0627a327cd774, points: 0, grade: true, solution: false -->
```python
ans_thirteen = answer_thirteen()
assert type(ans_thirteen) == nx.Graph , "Your return type should be a Graph object"

```

### Question 14

What is the transitivity and average clustering coefficient of graph G_un?

*This function should return a tuple (transitivity, avg clutering).*     
*Note: DO NOT round up your answer.*

<!-- grade_id: cell-63e48b1eed1a1fb8, points: 0, grade: false, solution: true -->
```python
def answer_fourteen():
    graph = answer_thirteen()
    return nx.transitivity(graph), nx.average_clustering(graph)
```

```python
answer_fourteen()
```

Output:
```
(0.570111160700385, 0.6975272437231418)
```

<!-- grade_id: cell-41dda91202f58e7e, points: 1, grade: true, solution: false -->
```python
ans_fourteen = answer_fourteen()
assert type(ans_fourteen) == tuple, "Student answer must be a tuple"

```

```python

```